<!-- DEBUT problem {{q.title}} -->
<problem display_name="{{ q.title }}" {% if q.type == "MULTICHOICE" or q.type == "TRUEFALSE" or q.type == 'MULTIANSWER' %}max_attempts="1"{%endif%}>
    <legend>
        {{ q_text }}
    </legend>
    {% if q.type == "MULTICHOICE" or q.type == "TRUEFALSE" %}
    <multiplechoiceresponse>
        <choicegroup type="MultipleChoice">
            {% for a in q.answers %}
            <choice correct="{{ 'true' if a.is_right else 'false'}}">{{a.answer_text|striptags}}
                <choicehint>{{a.feedback|striptags}}</choicehint>
            </choice>
            {% endfor %}
        </choicegroup>
    </multiplechoiceresponse>
    {% elif q.type == 'MULTIANSWER' %}
    <choiceresponse partial_credit="EDC">
        <checkboxgroup>
            {% for a in q.answers %}
            {% set correctness = 'true' if a.credit|float > 0.0 else 'false' %}
            <choice correct="{{correctness}}">
                {{a.answer_text|striptags}}
                {% if a.feedback|length > 1 %}
                <choicehint selected="true">
                    {{a.feedback|striptags}}
                </choicehint>
                {% endif %}
            </choice>
            {% endfor %}
        </checkboxgroup>
    </choiceresponse>
    {% elif q.type == 'ESSAY' %}
    {% include 'edx_problem_essay_script.xml' %}
    <customresponse cfn="checkAnswerEssay">
    	<textline size="40" correct_answer="" label="Problem Text"/>
    </customresponse> 
    {% endif %}
    {% if q.global_feedback|length > 1 %}
    <solution>
        <div class="detailed-solution">{{ q.global_feedback }}</div>
    </solution>
    {% endif %}
</problem>
<!-- FIN problem {{q.title}} -->
{#<demandhint>        <hint>A fruit is the fertilized ovary from a flower.</hint>        <hint>A fruit contains seeds of the plant.</hint>    </demandhint> #}